```python
# Initialize Otter
import otter
grader = otter.Notebook("hw02.ipynb")
```

## Homework 02

This is the second homework.

```python
import pandas as pd
import numpy as np
```

Here is a dataframe named `df` in which the `fwd` column contains the integers `0` through `10`. The `rev` column descends from `0` to `-10`.

```python
df = pd.DataFrame({
    'fwd': np.arange(10),
    'rev': np.arange(0, -10, -1)
})
df
```

Your job is to sort the rows of `df` using the values in the `rev` column. Assign the result to `sortdf`.

```python
sortdf = ...
```

```python
grader.check("1_create_df")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for 1_create_df (points: 2)
```
>>> try:
...     assert((sortdf['fwd'] == df['fwd'].iloc[::-1]).all())
...     assert((sortdf['rev'] == df['rev'].iloc[::-1]).all())
... except AssertionError:
...     print("sortdf is not properly sorted.")
```
