```python
# Initialize Otter
import otter
grader = otter.Notebook("hw01.ipynb")
```

# Homework 01

This is the first homework.

```python
import pandas as pd
```

Create an empty dataframe named `empty`.

```python
empty = ...
```

```python
grader.check("1_create_df")
```

Create a dataframe of two columns named `t1` and `t2`. The `t1` values must be `0.0` and `0.1`, and the `t2` values must be `0.1` and `0.2`. Assign the dataframe to a variable named `df`.

```python
# Your solution here
...
```

```python
grader.check("2_create_2col_df")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for 1_create_df (points: 2)
```
>>> assert(type(empty) == pd.core.frame.DataFrame)
```
```
>>> assert(len(empty) == 0)
```


Tests for 2_create_2col_df (points: 10)
```
>>> assert(type(df) == pd.core.frame.DataFrame)
```
```
>>> assert(set(df.columns) == set(['t1', 't2']))
```
```
>>> assert((df['t1'] == [0.0, 0.1]).all())
```
```
>>> assert((df['t2'] == [0.1, 0.2]).all())
```
```
>>> assert(len(df) == 2)
```
